# <p style="text-align: center;"> TP noté R2.07 <br> Quelques algorithmes sur les graphes
</p>

## <span style="color:red">  **Consignes générales :** </span>
1. Lire attentivement le sujet
2. les structures de données ainsi que  les entrées et sorties des méthodes demandées (et leur type) sont **imposées**. Il ne faut pas les modifier.
3. Pour qu'une cellule de code soit prise en compte elle doit êtr
4. e executée.
5. Si vous faites une modification dans une cellule, il faut "l'executer" pour qu'elle soit prise en compte.
6. Des tests de validation de certaines méthodes sont proposés, il est nécessaire de les passer mais pas suffisant !
7. Vous pouvez utiliser, en les adaptant,  les  méthodes obtenues durant les séances de tp. 


```python
from collections import deque
import random
import copy
g1 = {"A" :{"C"},
           "B" : {"C", "E"},
           "C" : {"A", "B", "D", "E"},
           "D" : {"C"},
           "E" : {"C", "B"},
           "F" : set()
         }
g2 ={"A" : set()}
g3 = {
"A":{"B":3,"C":4, "D":7},
"B":{"C":1,"F":5},
"C":{"F":6,"D":2},
"D":{"E":3, "G":6},
"E":{"G":3, "H":4},
"F":{"E":1, "H":8},
"G":{"H":2},
"H":{"G":2}
}
g4 = {
    "A": {"B" :3},
    "B": {"C":1, "F":4},
    "C": {"D":6, "E":2},
    "D": dict(),
    "E": {"F":4},
    "F": dict()
}
```

```python
class Graphe(object):
    
    def __init__(self, graphe_dict=None):
        """ initialise un objet graphe.
	    Si aucun dictionnaire n'est
	    créé ou donné, on en utilisera un 
	    vide
        """
        if graphe_dict == None:
            graphe_dict = dict()
        self._graphe_dict = graphe_dict

    def aretes(self, sommet):
        """ retourne une liste de toutes les aretes d'un sommet"""
        return self._graphe_dict[sommet]

    def all_sommets(self):
        """ retourne tous les sommets du graphe """
        return set(self._graphe_dict.keys())

    def all_aretes(self):
        """ retourne toutes les aretes du graphe """
        return self.__list_aretes()

    def add_sommet(self, sommet):
        """ Si le "sommet" n'set pas déjà présent
	dans le graphe, on rajoute au dictionnaire 
	une clé "sommet" avec une liste vide pour valeur. 
	Sinon on ne fait rien.
        """
        if sommet not in self._graphe_dict:
            self._graphe_dict[sommet] = []

    def add_arete(self, arete):
        """ l'arete est de  type set, tuple ou list;
            Entre deux sommets il peut y avoir plus
	    d'une arete (multi-graphe)
        """
        arete = set(arete)
        arete1, arete2 = tuple(arete)
        for x, y in [(arete1, arete2), (arete2, arete1)]:
            if x in self._graphe_dict:
                self._graphe_dict[x].add(y)
            else:
                self._graphe_dict[x] = {y}

    def __list_aretes(self):
        """ Methode privée pour récupérer les aretes. 
	    Une arete est un ensemble (set)
            avec un (boucle) ou deux sommets.
        """
        aretes = []
        for sommet in self._graphe_dict:
            for voisin in self._graphe_dict[sommet]:
                if ({voisin, sommet})  not in aretes:
                    aretes.append({sommet, voisin})
        return aretes
    
    def __iter__(self):
        self._iter_obj = iter(self._graphe_dict)
        return self._iter_obj

    def __next__(self):
        """ Pour itérer sur les sommets du graphe """
        return next(self._iter_obj)

    def __str__(self):
        res = "sommets: "
        for k in self._graphe_dict.keys():
            res += str(k) + " "
        res += "\naretes: "
        for arete in self.__list_aretes():
            res += str(arete) + " "
        return res
```

## Ecrire la méthode estcyclique, décrite dans l'algorithme ci dessous, qui permet de vérifier si un graphe est **acyclique**.


![Cycle](AlgoCyclique.png)

Si le sommet passé en paramètre n'existe pas, on renvoie une exception. On utlisera la méthode "raise" et  ValueError().

<!-- grade_id: cell-9700bf7fa8adb695, points: 0, grade: false, solution: true -->
```python
class Graphe2(Graphe):
    def estcyclique(self,s=None,TAG=[]):
        """ s est un sommet quelconque, donné par l'utilisateur 
        ou choisi aléatoirement parmi la liste des sommets du graphe
        """
        
        return C
```

```python
"""Vérifiez vos résultats sur les tests suivants"""
G=Graphe2(g1)
assert G.estcyclique("A") == "Il y a un cycle"
assert G.estcyclique("E") == "Il y a un cycle"

```

<!-- grade_id: cell-065a64a25b12eb9c, points: 1, grade: true, solution: false -->
```python
"""Vérifiez les entrées"""
try:
    G.estcyclique("Z") 
except ValueError:
    pass
else:
    raise AssertionError("n'a pas retourné d'erreur")
```

## Ecrire la méthode estconnexe, décrite ci-dessous, et qui permet de tester si un graphe est connexe

![Connexe](grapheconnexe.png)

<!-- grade_id: cell-4248824b00cbc543, points: 0, grade: false, solution: true -->
```python
class Graphe2(Graphe2):
    def estconnexe(self,s):
        """ s est un sommet quelconque, donné par l'utilisateur 
        Il faudra renvoyer une exception si le sommet n'existe pas
        """
        
        return "faux"
```

<!-- grade_id: cell-911ea2a527ab4f07, points: 1, grade: true, solution: false -->
```python
"""Vérifiez vos résultats sur les tests suivants"""
G=Graphe2(g1)
assert G.estconnexe("A") == "faux"
assert G.estconnexe("E") == "faux"

```

## Modifier la méthode **estconnexe** pour qu'elle retourne "vrai" si le graphe est connexe et,  sinon, le tuple ("faux",n) où n est le nombre de composantes connexes du graphe.

<!-- grade_id: cell-2e03fc6b59e99f40, points: 0, grade: false, solution: true -->
```python
class Graphe2(Graphe2):
    def estconnexe2(self,s):
        """ s est un sommet quelconque, donné par l'utilisateur 
        Il faudra renvoyer une exception si le sommet n'existe pas
        """
      
        return ("faux",n)
```

<!-- grade_id: cell-dfe93da380f98b32, points: 0, grade: true, solution: false -->
```python
"""Vérifiez vos résultats sur les tests suivants"""
G=Graphe2(g1)
assert G.estconnexe2("A") == ("faux",2)
assert G.estconnexe2("C") == ("faux",2)

```

## Modifier la méthode précédente pour qu"elle pour qu'elle retourne "vrai" si le graphe est connexe et,  sinon, <br> le tuple (("faux",n), L) où n est le nombre de composantes connexes du graphe et L une liste des composantes connexes (une liste de listes donc).

```python
class Graphe2(Graphe2):
    def estconnexe3(self,s):
        """ s est un sommet quelconque, donné par l'utilisateur 
        Il faudra renvoyer une exception si le sommet n'existe pas
        L=[[]]
        """
      
        return (("faux",n), L)
```

```python
"""Vérifiez vos résultats sur les tests suivants"""
G=Graphe2(g1)
assert G.estconnexe2("A") == (("faux",2), [["A","B","C","D","E"],["F"]])
assert G.estconnexe2("C") == (("faux",2), [["A","B","C","D","E"],["F"]])
```

## Dikjstra, plus long chemin

On souhaite ecrire une méthode qui renvoie le plus long chemin à partir d'un sommet donné.<br>
Pour cela on adaptera l'algorithme de Dikjstra pour des graphes  particuliers : **les graphes orientés acycliques pondérés positivement** 

## Ecrire, dans un premier temps, la méthode ci-dessous qui met en oeuvre l'algorithme de Dijkstra classique (D est le tableau des distances et P le dictionnaire des parents, comme vu en tp).

<!-- grade_id: cell-b992084d0ba8f7f3, points: 0, grade: false, solution: true -->
```python
class Graphe2(Graphe2):
    def dijkstra(self,s):
      
    return D,P
```

## Ecire la méthode ci dessous qui permet de donner le plus court chemin entre deux sommets, s et d, donnés.<br>Le chemin sera donné sous forme d'une liste dans le bon ordre.

<!-- grade_id: cell-0d85dbb2a97e06ea, points: 0, grade: false, solution: true -->
```python
class Graphe2(Graphe2):
    def chemindijkstra(self,s,d):
        """ Cette méthode retourne "pas de chemin" ou un tuple contenant
        le plus court chemin entre s et d sous forme d'une liste, L, dans le bon ordre
        ainsi que la longueur,l, du chemin trouvé. Le chemin d'un sommet 
        à lui mêeme est de longueur nulle
        Comme précédemment, on gèrera une exception si s ou d n'appartiennent pas au graphe 
        """
        
        return (L,l)
```

<!-- grade_id: cell-c530f17ffb5400d9, points: 3, grade: true, solution: false -->
```python
"""Vérifiez vos résultats sur les tests suivants"""
G3=Graphe2(g3)
assert G3.chemindijkstra("A","E") == (["A","C","D","E"],9)
assert G3.chemindijkstra("E","E") == (["E"],0)

```

On souhaite dans cette partie adapter l'algorithme de Dijkstra pour trouver le plus long chemin entre deux sommets d'un graphe valué orienté.
Pour cela :
1. Si le graphe est acyclique et pondéré positivement sinon on retourne "impossible"
2. Sinon, on utilise l'algorithme de Dijkstra sur le graphe mais en changeant toutes les pondérations en leurs opposées.
3. On retournera alors deux tuples : $(L_1,l_1)$ et $(L_2,l_2)$ où $L_1, L_2, l_1$ et $l_2$ sont respectivement les listes correspondant  au plus court, plus long chemin  (dans le bon ordre) et les longueurs de ces deux chemins
4. Comme précédemment si un des sommets n'existe pas on générera une exception. 

<!-- grade_id: cell-ce7d928e5bce5ed1, points: 0, grade: false, solution: true -->
```python
class Graphe2(Graphe2):
    def cheminsPLC(self,s,d):
        """ Cette méthode retourne "impossible" (si il y a un cycle ou si il y a une pondération négative),
        "pas de chemin" ou deux tuples contenant
        le plus court chemin entre s et d sous forme d'une liste, L1, dans le bon ordre
        ainsi que la longueur,l1, du chemin trouvé, le plus long chemin entre s et d sous forme d'une liste, 
        L2, dans le bon ordre ainsi que la longueur,l2, du chemin trouvé, Le chemin d'un sommet 
        à lui mêeme est de longueur nulle
        Comme précédemment, on gèrera une exception si s ou d n'appartiennent pas au graphe 
        """
 
       
     
        return (L1,l1), (L2,l2)
```